## Autograded Notebook (Canvas & CodeGrade)

This notebook will be automatically graded. It is designed to test your answers and award points for the correct answers. Following the instructions for each Task carefully.

### Instructions

* **Download this notebook** as you would any other ipynb file
* **Upload** to Google Colab or work locally (if you have that set-up)
* **Delete `raise NotImplementedError()`**
* Write your code in the `# YOUR CODE HERE` space
* **Execute** the Test cells that contain `assert` statements - these help you check your work (others contain hidden tests that will be checked when you submit through Canvas)
* **Save** your notebook when you are finished
* **Download** as a `ipynb` file (if working in Colab)
* **Upload** your complete notebook to Canvas (there will be additional instructions in Slack and/or Canvas)

# Data Science - Unit 1 Sprint 1 Module 1

## Exploratory Data Analysis

### Module Learning Objectives


- Load a CSV dataset from a URL using `pandas.read_csv`
- Load a CSV dataset from a local file using `pandas.read_csv
- Use basic Pandas functions for Exploratory Data Analysis-EDA
- Describe and discriminate between basic data types such as categorical, quantitative, continuous, discrete, ordinal, nominal and identifier

### Notebook points: 12

## Autograded Module Projects

Welcome to the first Module Project of Unit 1! You will complete a project (sometimes also referred to as an assignment) after the Guided Project for each Module. There will be four Module Projects per Sprint and each project is designed to provide you with an opportunity to practice what you learned in the Canvas Warm-up material and the Guided Project with your Instructors.

Throughout Unit 1, the Module Projects and the Sprint Challenges are *autograded*. You will complete your work in a Jupyter/Python notebook (the files that end with `.ipynb`) and then upload your completed notebook to Canvas and submit for grading. This autograding process will check your answers and provide more information about the errors in your notebook. You'll receive a score when the testing is complete.

So, let's get started! If you are reading this notebook, then you have the correct autograded version.

## Introduction

For this module, we learned how to use some of the most common tools for exploring our data. We'll continue to practice our `pandas` skills with a new dataset. 

## Dataset Description

Explore the University of California - Irvine Adult Dataset

**Task 1** - Load a dataset via its URL

* Create a Python list named `column_headers` with the following items: `'age', 'workclass', 'fnlwgt', 'education', 'education-num', 'marital-status', 'occupation', 'relationship', 'race', 'sex', 'capital-gain', 'capital-loss', 'hours-per-week','native-country', 'income'`.  Lists have the form `listname = ['element 1', 'element 2', element 3']`
* Load the file from the URL provided below using the `column_headers` list you created to name your columns: set the `pd.read_csv()` to `column_headers` using `names =`.  
* Name the DataFrame `adult`

<!-- grade_id: cell-8a338d1f29fadeee, points: 0, grade: false, solution: true -->
```python
# Task 1
import pandas as pd
import numpy as np

# URL for the dataset
data_url = 'https://archive.ics.uci.edu/ml/machine-learning-databases/adult/adult.data'

# YOUR CODE HERE
raise NotImplementedError()

# Print out your DataFrame
adult.head()
```

**Task 1 Test**

<!-- grade_id: cell-6593303164ade3dd, points: 1, grade: true, solution: false -->
```python
# Task 1 - Test

# These tests are for you to check your work before submitting
assert isinstance(column_headers, list), 'Have you created a list of header names?'
assert len(column_headers) == 15, 'Did you include the correct number of items in the column_headers list?'
assert isinstance(adult, pd.DataFrame), 'Have you created a DataFrame named adult?'

# These tests will be completed when you submit your notebook for autograding
```

**Task 2** - Look at the first and last rows of the DataFrame

* Assign the first **ten (10)** rows of the `adult` DataFrame to `adult_head`
* Assign the last **ten (10)** rows of the `adult` DataFrame to `adult_tail`

<!-- grade_id: cell-a9fb33ec1c6025a8, points: 0, grade: false, solution: true -->
```python
# Task 2

# YOUR CODE HERE
raise NotImplementedError()

# Optional: print out adult_head and adult_tail
```

**Task 2 Test**

<!-- grade_id: cell-db111c4de2805de2, points: 1, grade: true, solution: false -->
```python
# Task 2 - Test

# These tests are for you to check your work before submitting
assert isinstance(adult_head, pd.DataFrame), 'Have you created a DataFrame named adult_head?'
assert isinstance(adult_tail, pd.DataFrame), 'Have you created a DataFrame named adult_tail?'

# These tests will be completed when you submit your notebook for autograding
```

**Task 3** - Variable data types

For your `adult` DataFrame, determine the data types for the variables.

* Count the number of `int64` variable types and assign to `number_int64` (your value should be an integer)
* Count the number of `object` variable types and assign to `number_object` (your value should be an integer)

There are different ways to find the data types of the columns in a DataFrame. All you need to do for this task is to count the number of each data type and assign to the corresponding variable.

<!-- grade_id: cell-340508752abebcbf, points: 0, grade: false, solution: true -->
```python
# Task 3

# YOUR CODE HERE
raise NotImplementedError()
```

**Task 3 Test**

<!-- grade_id: cell-05ba0ce5673511d6, points: 1, grade: true, solution: false -->
```python
# Task 3 - Test

# Hidden tests - you will see the results when you submit to Canvas
```

**Task 4** - DataFrame dimensions

* Find the dimensions of your DataFrame and assign result to the variable `adult_dimension`. Your variable should be a *tuple* and the row dimension should be listed first.

Hint: A tuple looks like this: (1, 2) - a tuple is a collection which is ordered and unchangeable. You can use the function `type(your_variable)` to print the variable type.

Hint 2: The `shape` method returns a tuple - convenient!

<!-- grade_id: cell-d70ccb5657fd3d7b, points: 0, grade: false, solution: true -->
```python
# Task 4

# YOUR CODE HERE
raise NotImplementedError()

# Print out the shape
print(adult_dimension)
```

**Task 4 Test**

<!-- grade_id: cell-6826bc35a7f56e9d, points: 1, grade: true, solution: false -->
```python
# Task 4 - Test

assert isinstance(adult_dimension, tuple), 'Have you created a tuple named adult_dimension?'

# Hidden tests - you will see the results when you submit to Canvas
```

**Task 5** - Missing values

Are there any missing values in the dataset? Let's check!

* Check for missing values using `.isnull().sum()`
* Count the number of missing values and assign the value to the variable `adult_missing`. Your variable should be an integer.

<!-- grade_id: cell-061626ea17d6ee51, points: 0, grade: false, solution: true -->
```python
# Task 5

# YOUR CODE HERE
raise NotImplementedError()
```

**Task 5 Test**

<!-- grade_id: cell-fa29b69831f8efb7, points: 1, grade: true, solution: false -->
```python
# Task 5 - Test

# Hidden tests - you will see the results when you submit to Canvas
```

**Task 6** - Viewing DataFrame statistics

Let's look at some of the *summary statistics* of our dataset. We can use the .describe() function in order to see the summary statistics of the numeric columns.

Look at the statistics for the `adult` DataFrame and then complete the following two tasks:

* Find the value for the mean `age` and assign it to the variable `mean_age` (your value should be a float and defined to two decimal places)
* Find the standard deviation (std) for the `hours-per-week` variable and assign it to `std_hpw` (your value should be a float and defined to two decimal places)

Are there any values shown that might be a code for missing data?  We'll learn how to change the values in a DataFrame in the next module.

<!-- grade_id: cell-379c9d6a1454cb65, points: 0, grade: false, solution: true -->
```python
# Task 6

# YOUR CODE HERE
raise NotImplementedError()
```

**Task 6 Test**

<!-- grade_id: cell-26b21a59dd35ddb1, points: 1, grade: true, solution: false -->
```python
# Task 6 - Test

# Hidden tests - you will see the results when you submit to Canvas
```

**Task 7** - Summary statistics for non-numeric columns

We have some columns in this dataset that are non-numeric or object columns. This is usually a string object. Let's use the `describe()` function again but include the argument `exclude='number'` to exclude the numeric columns.

Using the results of `describe(exclude='number')`, complete the following two tasks:

* Find the number of unique `education` values and assign it to `unique_edu` (your value should be an integer)
* Find the number of times the most frequent observation for `income` occurs and assign it to `freq_income` (your value should be an integer) *(Note: this is not the income value itself, just how many times it occurs)*

<!-- grade_id: cell-11915c7656419ee3, points: 0, grade: false, solution: true -->
```python
# Task 7

# YOUR CODE HERE
raise NotImplementedError()
```

```python
adult.describe(exclude='number')
adult['income'].value_counts()
```

**Task 7 Test**

<!-- grade_id: cell-00ef7970684097e5, points: 1, grade: true, solution: false -->
```python
# Task 7 - Test

# Hidden tests - you will see the results when you submit to Canvas
```

**Task 8** - Finding value counts

Let's look more specifically at the `relationship` column and perform a count of the number of observations for each category. We can see how many categories we have by using the `.unique()` method on the column. Then, we can use `.value_counts()` to count the number of observations in each category.

* View the unique values in the `relationship` column with `.unique()`
* View the number of observations for each value with `.value_counts()`
* Find the counts for `Other-relative` and assign to the variable `adult_other_rel`

<!-- grade_id: cell-2498319aac6c6e38, points: 0, grade: false, solution: true -->
```python
# Task 8

# YOUR CODE HERE
raise NotImplementedError()
```

**Task 8 Test**

<!-- grade_id: cell-ee00a51b05f5cc8e, points: 1, grade: true, solution: false -->
```python
# Task 8 - Test

# Hidden tests - you will see the results when you submit to Canvas
```

**Task 9** - Create a Series

A column of a DataFrame is a pandas Series. Using the `adult` DataFrame, create a Series from the `occupation` column.

* Create a Series from the `occupation` column and name it `adult_occup`

<!-- grade_id: cell-d1bb196ab00c4bd3, points: 0, grade: false, solution: true -->
```python
# Task 9

# YOUR CODE HERE
raise NotImplementedError()
```

**Task 9 Test**

<!-- grade_id: cell-e43fc22d03a71561, points: 1, grade: true, solution: false -->
```python
# Task 9 - Test

# These tests are for you to check your work before submitting
assert isinstance(adult_occup, pd.Series), 'Have you created the correct Series?'

# Hidden tests - you will see the results when you submit to Canvas
```

**Task 10** - Practice with a new dataset

Let's use some of what we've learned so far to load in a new dataset and answer some questions. For now, we're going to read in the data from the website where it is stored. But, at the end of this project, you can review the instructions for uploading a file to Google Colab,

* Use [this link (right click to copy)](https://raw.githubusercontent.com/LambdaSchool/data-science-practice-datasets/main/unit_1/NASCAR/nascard.csv) as the URL and assign it to the variable `data_url2`
* Read in the data into a DataFrame called `nascar`; the CSV already includes a header

More information about this data can be found [here](https://github.com/BloomInstituteOfTechnology/data-science-practice-datasets/tree/main/unit_1/NASCAR).

<!-- grade_id: cell-04b98cad3de4cab4, points: 0, grade: false, solution: true -->
```python
# Task 10

# YOUR CODE HERE
raise NotImplementedError()
```

**Task 10 Test**

<!-- grade_id: cell-8b7550f8a360559b, points: 1, grade: true, solution: false -->
```python
# Task 10 - Test

# These tests are for you to check your work before submitting
assert isinstance(nascar, pd.DataFrame), 'Have you created the nascar DataFrame?'

# Hidden tests - you will see the results when you submit to Canvas
```

**Task 11** - Look at the header

Now, let's look at the data. We'll use the `.head()` method to view the first rows of the DataFrame

* Assign the output from `.head()` to a DataFrame called `nascar_head` (use the default values of the method)

<!-- grade_id: cell-01b2cebe2cab129a, points: 0, grade: false, solution: true -->
```python
# Task 11

# YOUR CODE HERE
raise NotImplementedError()
```

**Task 11 Test**

<!-- grade_id: cell-4abee0d24322116e, points: 1, grade: true, solution: false -->
```python
# Task 11 - Test

# These tests are for you to check your work before submitting
assert isinstance(nascar_head, pd.DataFrame), 'Have you created the nascar_head DataFrame?'

# Hidden tests - you will see the results when you submit to Canvas
```

**Task 12** - Dataset description and variable types

Using your `nascar` DataFrame, answer the following questions. Remember - you can use methods like `describe()` and `info()` to print out information about the variable types in your DataFrame. Remember, you can use `describe(exclude='number')` to view non-numeric variable types.

You can view more information about the data [here](http://users.stat.ufl.edu/~winner/data/nascard.txt).

Select the letter indicating which one of the following statements about the features in `nascar` is incorrect?  For example, if statement B is incorrect, you'll type your answer in the code block as `answer = 'B'`


A: Each entry for `driver` is a text string

B: `carMake` is an identifer variable


C: `carsRace` is a quantitative, discrete variable


D: `prize` is a quantitative, continuous variable

<!-- grade_id: cell-4452c94a7f85fa40, points: 0, grade: false, solution: true -->
```python
# Task 12

# Ignore the YOUR CODE HERE for this cell
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-2673e645b1df248f, points: 1, grade: true, solution: false -->
```python
# Task 12 - Test

# Hidden tests - you will see the results when you submit to Canvas
```

**Extra Practice!!**

Now we're going to practice loading a dataset that you have saved on your computer. Since we're working in a notebook on Google Colab, you'll need to upload your file to the notebook runtime in order to read in the data.

* Use [this link](https://raw.githubusercontent.com/LambdaSchool/data-science-practice-datasets/main/unit_1/NASCAR/nascard.csv) to download your dataset as a csv file to your computer. 
* Upload to Google Colab one of two ways:
    * using `from google.colab import files` then `uploaded = files.upload()`; read in the file with pd.read_csv(filename)
    * using the file upload feature in the left side panel of the Colab notebook; read in the file with pd.read_csv(filename)

More information about this data can be found [here](https://github.com/BloomInstituteOfTechnology/data-science-practice-datasets/tree/main/unit_1/NASCAR).

**DELETE or COMMENT OUT your Colab code**

```python
# Extra Practice

## DELETE OR COMMENT OUT YOUR COLAB CODE BEFORE SUBMITTING THE NOTEBOOK
```