<a href="https://colab.research.google.com/github/lucaanescu/DS-Unit-1-Sprint-1-Data-Wrangling-and-Storytelling/blob/master/LS_DS_Sprint1_AG.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

## Autograded Notebook (Canvas & CodeGrade)

This notebook will be automatically graded. It is designed to test your answers and award points for the correct answers. Following the instructions for each Task carefully.

### Instructions

* **Download this notebook** as you would any other ipynb file
* **Upload** to Google Colab or work locally (if you have that set-up)
* **Delete `raise NotImplementedError()`**
* Write your code in the `# YOUR CODE HERE` space
* **Execute** the Test cells that contain `assert` statements - these help you check your work (others contain hidden tests that will be checked when you submit through Canvas)
* **Save** your notebook when you are finished
* **Download** as a `ipynb` file (if working in Colab)
* **Upload** your complete notebook to Canvas (there will be additional instructions in Slack and/or Canvas)

# Lambda School Data Science - Unit 1 Sprint 1

## Sprint Challenge: Data Wrangling and Storytelling


## Use the following information to complete Tasks 1 - 12

### Notebook points total: 12

In this Sprint Challenge you will first "wrangle" some data from [Gapminder](https://www.gapminder.org/about-gapminder/), a Swedish non-profit co-founded by Hans Rosling. "Gapminder produces free teaching resources making the world understandable based on reliable statistics."
- [Cell phones (total), by country and year](https://raw.githubusercontent.com/open-numbers/ddf--gapminder--systema_globalis/master/countries-etc-datapoints/ddf--datapoints--cell_phones_total--by--geo--time.csv)
- [Population (total), by country and year](https://raw.githubusercontent.com/open-numbers/ddf--gapminder--systema_globalis/master/countries-etc-datapoints/ddf--datapoints--population_total--by--geo--time.csv)
- [Geo country codes](https://github.com/open-numbers/ddf--gapminder--systema_globalis/blob/master/ddf--entities--geo--country.csv)

These two links have everything you need to successfully complete the first part of this sprint challenge.
- [Pandas documentation: Working with Text Data](https://pandas.pydata.org/pandas-docs/stable/text.html) (one question)
- [Pandas Cheat Sheet](https://github.com/pandas-dev/pandas/blob/master/doc/cheatsheet/Pandas_Cheat_Sheet.pdf) (everything else)


**Task 1** - Load and print the cell phone data.  Pandas and numpy import statements have been included for you.

* load your CSV file found at `cell_phones_url` into a DataFrame named `cell_phones`
* print the top 5 records of `cell_phones`


<!-- grade_id: cell-6f16afb9d271f949, points: 0, grade: false, solution: true -->
```python
# Task 1

# Imports 
import pandas as pd
import numpy as np

cell_phones_url = 'https://raw.githubusercontent.com/LambdaSchool/data-science-practice-datasets/main/unit_1/Cell__Phones/cell_phones.csv'

# Load the dataframe and print the top 5 rows

# YOUR CODE HERE
raise NotImplementedError()

```

**Task 1 Test**

<!-- grade_id: cell-226f7bf8e9ea24f9, points: 1, grade: true, solution: false -->
```python
# Task 1 - Test

assert isinstance(cell_phones, pd.DataFrame), 'Have you created a DataFrame named `cell_phones`?'
assert len(cell_phones) == 9574

```

**Task 2** - Load and print the population data.  

* load the CSV file found at `population_url` into a DataFrame named `population`
* print the top 5 records of `population`


<!-- grade_id: cell-561c2d59728188a9, points: 0, grade: false, solution: true -->
```python
# Task 2

population_url = 'https://raw.githubusercontent.com/LambdaSchool/data-science-practice-datasets/main/unit_1/Population/population.csv'

# Load the dataframe and print the first 5 records

# YOUR CODE HERE
raise NotImplementedError()
```

**Task 2 Test**

<!-- grade_id: cell-59d01cd695becd74, points: 1, grade: true, solution: false -->
```python
# Task 2 - Test

assert isinstance(population, pd.DataFrame), 'Have you created a DataFrame named `population`?'
assert len(population) == 59297

```

**Task 3** - Load and print the geo country codes data.  

* load the CSV file found at `geo_codes_url` into a DataFrame named `geo_codes`
* print the top 5 records of `geo_codes`


<!-- grade_id: cell-eb4d290384535503, points: 0, grade: false, solution: true -->
```python
# Task 3

geo_codes_url = 'https://raw.githubusercontent.com/LambdaSchool/data-science-practice-datasets/main/unit_1/GEO_codes/geo_country_codes.csv'

# Load the dataframe and print out the first 5 records

# YOUR CODE HERE
raise NotImplementedError()
```

**Task 3 Test**

<!-- grade_id: cell-39240405659c0c19, points: 1, grade: true, solution: false -->
```python
# Task 3 - Test

assert geo_codes is not None, 'Have you created a DataFrame named `geo_codes`?'
assert len(geo_codes) == 273

```

**Task 4** - Check for missing values

Let's check for missing values in each of these DataFrames: `cell_phones`, `population` and `geo_codes`

* Check for missing values in the following DataFrames:
    * assign the total number of missing values in `cell_phones` to the variable `cell_phones_missing`
    * assign the total number of missing values in `population` to the variable `population_missing`
    * assign the total number of missing values in `geo_codes` to the variable `geo_codes_missing` (Hint: you will need to do a sum of a sum here - `.sum().sum()`)

<!-- grade_id: cell-9426cd5765574e07, points: 0, grade: false, solution: true -->
```python
# Task 4

# Check for missing data in each of the DataFrames

# YOUR CODE HERE
raise NotImplementedError()
```

**Task 4 Test**

<!-- grade_id: cell-cf6ab3b4b1e8afc1, points: 1, grade: true, solution: false -->
```python
# Task 4 - Test

if geo_codes_missing == 21: print('ERROR: Make sure to use a sum of a sum for the missing geo codes!')  

# Hidden tests - you will see the results when you submit to Canvas
```

**Task 5** - Merge the `cell_phones` and `population` DataFrames.

* Merge the `cell_phones` and `population` dataframes with an **inner** merge on `geo` and `time`
* Call the resulting dataframe `cell_phone_population`

<!-- grade_id: cell-decaebaa844aa3a5, points: 0, grade: false, solution: true -->
```python
# Task 5

# Merge the cell_phones and population dataframes

# YOUR CODE HERE
raise NotImplementedError()
```

**Task 5 Test**

<!-- grade_id: cell-dd2473ea91f15f30, points: 1, grade: true, solution: false -->
```python
# Task 5 - Test

assert cell_phone_population is not None, 'Have you merged created a DataFrame named cell_phone_population?'
assert len(cell_phone_population) == 8930

```

**Task 6** - Merge the `cell_phone_population` and `geo_codes` DataFrames

* Merge the `cell_phone_population` and `geo_codes` DataFrames with an inner merge on `geo`
* **Only merge in the `country` and `geo` columns from `geo_codes`**
* Call the resulting DataFrame `geo_cell_phone_population`


<!-- grade_id: cell-1ce5a2360ee6fd20, points: 0, grade: false, solution: true -->
```python
# Task 6

# Merge the cell_phone_population and geo_codes dataframes
# Only include the country and geo columns from geo_codes

# YOUR CODE HERE
raise NotImplementedError()
```

```python
geo_cell_phone_population.head()
```

**Task 6 Test**

<!-- grade_id: cell-764d5b72ae382339, points: 1, grade: true, solution: false -->
```python
# Task 6 - Test

assert geo_cell_phone_population is not None, 'Have you created a DataFrame named geo_cell_phone_population?
assert len(geo_cell_phone_population) == 8930

```

**Task 7** - Calculate the number of cell phones per person.

* Use the `cell_phones_total` and `population_total` columns to calculate the number of cell phones per person for each country and year.
* Call this new feature (column) `phones_per_person` and add it to the `geo_cell_phone_population` DataFrame (you'll be adding the column to the DataFrame).


<!-- grade_id: cell-bd0bce2920604643, points: 0, grade: false, solution: true -->
```python
# Task 7

# Calculate the number of cell phones per person for each country and year.

# YOUR CODE HERE
raise NotImplementedError()
```

**Task 7 Test**

<!-- grade_id: cell-45c955c43e471400, points: 1, grade: true, solution: false -->
```python
# Task 7 - Test

# Hidden tests - you will see the results when you submit to Canvas
```

**Task 8** - Identify the number of cell phones per person in the US in 2017

* Write a line of code that will create a one-row subset of `geo_cell_phone_population` with data on cell phone ownership in the USA for the year 2017.
* Call this subset DataFrame `US_2017`.
* Print `US_2017`.

<!-- grade_id: cell-665e83d11e594d90, points: 0, grade: false, solution: true -->
```python
# Task 8

# Determine the number of cell phones per person in the US in 2017

# YOUR CODE HERE
raise NotImplementedError()

# View the DataFrame
US_2017
```

**Task 8 Test**

<!-- grade_id: cell-ea08fdda80cf9731, points: 1, grade: true, solution: false -->
```python
# Task 8 - Test

# Hidden tests - you will see the results when you submit to Canvas
```

**Task 9** - Describe the numeric variables in `geo_cell_phone_population`

* Calculate the summary statistics for the quantitative variables in `geo_cell_phone_population` using `.describe()`.
* Find the mean value for `phones_per_person` and assign it to the variable `mean_phones`. Define your value out to two decimal points.


<!-- grade_id: cell-181c9805c52dfda8, points: 0, grade: false, solution: true -->
```python
# Task 9

# Calculate the summary statistics for the quantitative variables in geo_cell_phone_population using .describe()

# YOUR CODE HERE
raise NotImplementedError()
```

**Task 9 Test**

<!-- grade_id: cell-0a0cfb1ac7acc279, points: 1, grade: true, solution: false -->
```python
# Task 9 - Test

# Hidden tests - you will see the results when you submit to Canvas
```

**Task 10** - Describe the categorical variables in `geo_cell_phone_population`

* Calculate the summary statistics for the categorical variables in `geo_cell_phone_population` using `.describe(exclude='number')`.
* Using these results, find the number of unique countries and assign it to the variable `unique_country`. Your value should be an integer.

<!-- grade_id: cell-07de5adb6a8b43bf, points: 0, grade: false, solution: true -->
```python
# Task 10

# Calculate the summary statistics in geo_cell_phone_population using .describe(exclude='number')

# YOUR CODE HERE
raise NotImplementedError()
```

**Task 10 Test**

<!-- grade_id: cell-7672dd0b5f927ce7, points: 1, grade: true, solution: false -->
```python
# Task 10 - Test

# Hidden tests - you will see the results when you submit to Canvas
```

**Task 11** - Subset the DataFrame for 2017

* Create a new dataframe called `df2017` that includes **only** records from `geo_cell_phone_population` that ocurred in 2017.

<!-- grade_id: cell-f3fa17e15b5174ac, points: 0, grade: false, solution: true -->
```python
# Task 11

# Create a new dataframe called df2017 that includes only records from geo_cell_phone_population that ocurred in 2017.

# YOUR CODE HERE
raise NotImplementedError()
```

**Task 11 Test**

<!-- grade_id: cell-b19a6956b5dddadb, points: 1, grade: true, solution: false -->
```python
# Task 11 - Test

# Hidden tests - you will see the results when you submit to Canvas
```

**Task 12** - Identify the five countries with the most cell phones per person in 2017

* Sort the `df2017` DataFrame by `phones_per_person` in descending order and assign the result to `df2017_top`. Your new DataFrame should only have **five** rows (Hint: use `.head()` to return only five rows).
* Print the first 5 records of `df2017_top`.

<!-- grade_id: cell-c9b701bfe897edf5, points: 0, grade: false, solution: true -->
```python
# Task 12

# Sort the df2017 dataframe by phones_per_person in descending order
# Return only five (5) rows

# YOUR CODE HERE
raise NotImplementedError()

# View the df2017_top DataFrame
df2017_top
```

**Task 12 Test**

<!-- grade_id: cell-5f6fffc9db1b9492, points: 1, grade: true, solution: false -->
```python
# Task 12 - Test

assert df2017_top.shape == (5,6), 'Make sure you return only five rows'

```

**Task 13** - Explain why the figure below **cannot** be graphed as a pie chart.


```python
from IPython.display import display, Image
png = 'https://fivethirtyeight.com/wp-content/uploads/2014/04/hickey-ross-tags-1.png'
example = Image(png, width=500)
display(example)
```

**Task 13 Question** - Explain why the figure cannot be graphed as a pie chart.

*This task will not be autograded - but it is part of completing the challenge.*

YOUR ANSWER HERE

**Task 14** - Titanic dataset

Use the following Titanic DataFrame to complete Task 14 - execute the cell to load the dataset.

```python
# Load the Titanic dataset from its URL:
Titanic = pd.read_csv('https://raw.githubusercontent.com/LambdaSchool/data-science-practice-datasets/main/unit_1/Titanic/Titanic.csv')

Titanic.head()
```

**Task 14** - Create a visualization to show the distribution of **Parents/Children_Aboard**.

*This task will not be autograded - but it is part of completing the challenge.*

```python
### YOUR CODE HERE ###
```

Describe the distribution of Parents/Children_Aboard.

YOUR ANSWER HERE