<a href="https://colab.research.google.com/github/johhan27/DS-Unit-1-Sprint-1-Data-Wrangling-and-Storytelling/blob/master/LS_DS_113_Join_and_Reshape_Data_Assignment_AG.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

## Autograded Notebook (Canvas & CodeGrade)

This notebook will be automatically graded. It is designed to test your answers and award points for the correct answers. Following the instructions for each Task carefully.

### Instructions

* **Download this notebook** as you would any other ipynb file
* **Upload** to Google Colab or work locally (if you have that set-up)
* **Delete `raise NotImplementedError()`**
* Write your code in the `# YOUR CODE HERE` space
* **Execute** the Test cells that contain `assert` statements - these help you check your work (others contain hidden tests that will be checked when you submit through Canvas)
* **Save** your notebook when you are finished
* **Download** as a `ipynb` file (if working in Colab)
* **Upload** your complete notebook to Canvas (there will be additional instructions in Slack and/or Canvas)

# Lambda School Data Science - Unit 1 Sprint 1 Module 3

## Join and Reshape Data 

### Module Learning Objectives

* Concatenate data using the pandas concat method
* Merge data using pandas merge
* Define the concept of tidy data and describe the format
* Tidy data

### Helpful Resources

* [Pandas Cheat Sheet](https://github.com/pandas-dev/pandas/blob/master/doc/cheatsheet/Pandas_Cheat_Sheet.pdf)
* Python Data Science Handbook
    * [Chapter 3.6](https://jakevdp.github.io/PythonDataScienceHandbook/03.06-concat-and-append.html), Combining Datasets: Concat and Append
    * [Chapter 3.7](https://jakevdp.github.io/PythonDataScienceHandbook/03.07-merge-and-join.html), Combining Datasets: Merge and Join
    * [Chapter 3.8](https://jakevdp.github.io/PythonDataScienceHandbook/03.08-aggregation-and-grouping.html), Aggregation and Grouping
    * [Chapter 3.9](https://jakevdp.github.io/PythonDataScienceHandbook/03.09-pivot-tables.html), Pivot Tables

## Introduction

In this Module Project, we're going to focus on manipulating data: loading multiple datasets from different locations, concatenating, merging, and even the fancy pandas `groupby` method.

## Dataset Description

You can view the CSV files [here](https://github.com/LambdaSchool/data-science-practice-datasets/tree/main/unit_1/Instacart_small), before you load them into your notebook.

**Task 1** - Load CSV files

We're going to look at the Instacart data and load the files as we need them. The Pandas and numpy import statements have been provided for you as well as the URL for the dataset.

* Using the provided URL (`data_url1`) load the product CSV file into a DataFrame called `product_names`.
* Print the top five records (or rows) of `product_names`.

<!-- grade_id: cell-ef51cab74649e553, points: 0, grade: false, solution: true -->
```python
# Task 1

# Imports 
import pandas as pd
import numpy as np

# Dataset URL
data_url1 = 'https://raw.githubusercontent.com/LambdaSchool/data-science-practice-datasets/main/unit_1/Instacart_small/product_names_sm.csv'
# YOUR CODE HERE
product_names = pd.read_csv(data_url1)

# View the DataFrame
product_names.head()
```

Output:
```
product_id  ... department_id
0       45676  ...            11
1       35408  ...            16
2       37442  ...            16
3       41599  ...            11
4       21994  ...            11

[5 rows x 4 columns]
```

<!-- grade_id: cell-cf00b9647815499e, points: 1, grade: true, solution: false -->
```python
# Task 1 - Test

assert isinstance(product_names, pd.DataFrame), 'Have you created a DataFrame named `product_names`?'
assert len(product_names) == 25000

```

**Task 2** - DataFrame concatenation

Next, we're going to load two additional CSV files. The URLs have been provided so your task is to load the files into the appropriately named DataFrames and then to concatenate them into one DataFrame.

* Load the CSV file at `data_url2` into a DataFrame called `order_products_prior`.
* Load the CSV file at `data_url3` into a DataFrame called `order_products_train`.
* Concatenate `order_products_prior` and `order_products_train` into one DataFrame called `order_products`.

<!-- grade_id: cell-0f442142d6bc8c39, points: 0, grade: false, solution: true -->
```python
# Task 2

# Dataset URLs
data_url2 = 'https://raw.githubusercontent.com/LambdaSchool/data-science-practice-datasets/main/unit_1/Instacart_small/order_products_prior_sm.csv'
data_url3 = 'https://raw.githubusercontent.com/LambdaSchool/data-science-practice-datasets/main/unit_1/Instacart_small/order_products_train_sm.csv'

# YOUR CODE HERE
order_products_prior = pd.read_csv(data_url2)
order_products_train = pd.read_csv(data_url3)

order_products = pd.concat([order_products_prior, order_products_train], ignore_index=True)
# Print the head to view your results
order_products.tail()
```

Output:
```
order_id  product_id  add_to_cart_order  reordered
49995   1474983       21903                 21          0
49996    891325       26209                  5          1
49997    151280       13176                  6          1
49998   3250616       47719                  4          0
49999   2733689       17435                 19          0
```

<!-- grade_id: cell-3637e4ba19b79349, points: 1, grade: true, solution: false -->
```python
# Task 2 - Test

assert isinstance(order_products, pd.DataFrame), 'Have you created a DataFrame named `order_products`?'
assert len(order_products) == 50000
#24852 #
```

**Task 3** - Top five (5) products

We're going to create a list of the top five products - we've already determined what the top five are so now it's your turn to put these items in a list.

* Create a list called `five_products` with the following items: "Bag of Organic Bananas", "Organic Baby Spinach", "Organic Avocado", "Large Lemon", "Limes"

*Note: Make sure to create the list with exact spelling as shown above!*

```python
top_ten = order_products.groupby("product_id").count().sort_values(by=["order_id"], ascending=False).head(10)
top_ten 
```

Output:
```
order_id  add_to_cart_order  reordered
product_id                                        
24852            719                719        719
13176            561                561        561
21137            404                404        404
21903            360                360        360
47209            307                307        307
47626            283                283        283
47766            269                269        269
16797            220                220        220
27966            211                211        211
27845            201                201        201
```

```python
order_products["product_id"].value_counts().head(10) #alternative way to do the previous line 
```

Output:
```
24852    719
13176    561
21137    404
21903    360
47209    307
47626    283
47766    269
16797    220
27966    211
27845    201
Name: product_id, dtype: int64
```

```python
top_ten_list = list(top_ten.index)
```

```python
for name in top_ten_list:
  print (product_names.loc[product_names["product_id"] == name]["product_name"])
```

Output:
```
Series([], Name: product_name, dtype: object)
8384    Bag of Organic Bananas
Name: product_name, dtype: object
Series([], Name: product_name, dtype: object)
11316    Organic Baby Spinach
Name: product_name, dtype: object
Series([], Name: product_name, dtype: object)
6968    Large Lemon
Name: product_name, dtype: object
5991    Organic Avocado
Name: product_name, dtype: object
Series([], Name: product_name, dtype: object)
Series([], Name: product_name, dtype: object)
Series([], Name: product_name, dtype: object)
```

```python
#df.loc[df['channel'].isin(['sale','fullprice'])]
product_names.loc[product_names["product_id"].isin(top_ten_list)]
```

Output:
```
product_id            product_name  aisle_id  department_id
5991        47766         Organic Avocado        24              4
6968        47626             Large Lemon        24              4
8384        13176  Bag of Organic Bananas        24              4
11316       21903    Organic Baby Spinach       123              4
```

<!-- grade_id: cell-ad8e4664b33be149, points: 0, grade: false, solution: true -->
```python
# Task 3

# Top 5 products: five_products

# YOUR CODE HERE
five_products = ["Bag of Organic Bananas", "Organic Baby Spinach", "Organic Avocado", "Large Lemon", "Limes"]
```

<!-- grade_id: cell-0e8a931d7088b469, points: 1, grade: true, solution: false -->
```python
# Task 3 - Test

assert isinstance(five_products, list), 'Have you created the five_products list?'
assert len(five_products) == 5, "Make sure you created the list correctly"

# No hidden tests for this task
```

**Task 4** - DataFrame subset using merge

Create a subset of the `product_names` dataset that includes just the top **five** most ordered products. The `five_products` list has been converted to a DataFrame with the name `five`. Your task is to merge these DataFrames.

* Merge `product_names` and `five` so the resulting DataFrame only includes the 5 most ordered products.  Name that DataFrame `top_products`.

<!-- grade_id: cell-7818f18ec6921fb0, points: 0, grade: false, solution: true -->
```python
# Task 4

# Convert five_products to a DataFrame called five
five = pd.DataFrame(five_products, columns =['product_name']) 

# Create a subset of product_names with only the top 5 products.  
# Use your merge statement to include only the top 5 products

# YOUR CODE HERE
top_products = five.merge(product_names, how='left' , on='product_name')

# Print your DataFrame to view the merge results
top_products.head()
```

Output:
```
product_name  product_id  aisle_id  department_id
0  Bag of Organic Bananas       13176        24              4
1    Organic Baby Spinach       21903       123              4
2         Organic Avocado       47766        24              4
3             Large Lemon       47626        24              4
4                   Limes       26209        24              4
```

```python
product_names.loc[product_names["product_name"]=="Limes"]
```

Output:
```
product_id product_name  aisle_id  department_id
8939       26209        Limes        24              4
```

```python
order_products["product_id"].value_counts().loc[26209]
```

Output:
```
197
```

<!-- grade_id: cell-e329e3cdc129ba49, points: 1, grade: true, solution: false -->
```python
# Task 4 - Test
assert isinstance(top_products, pd.DataFrame), 'Have you created a DataFrame named `top_products`?'

```

**Task 5** - More DataFrame merging

Now we're going to merge the first DataFrame we created (`order_products`) with the one we just made above (`top_products`). This will result in a DataFrame with only the top five most ordered products.

* Merge `order_products` and `top_products` and name the resulting DataFrame `final`.
* Print the top 20 rows of `final` to view your result.

<!-- grade_id: cell-977f302dd89af1df, points: 0, grade: false, solution: true -->
```python
# Task 5

# Merge order_products and top_products 

# YOUR CODE HERE
final = order_products.merge(top_products, how="inner") #right would also work 

# Print your final DataFrame
final.head(20)
```

Output:
```
order_id  product_id  ...  aisle_id  department_id
0    2469765       26209  ...        24              4
1    2322104       26209  ...        24              4
2    1297655       26209  ...        24              4
3    1397335       26209  ...        24              4
4     651387       26209  ...        24              4
5    1338110       26209  ...        24              4
6    2587073       26209  ...        24              4
7      25158       26209  ...        24              4
8    2652386       26209  ...        24              4
9     440430       26209  ...        24              4
10   1080877       26209  ...        24              4
11    105005       26209  ...        24              4
12    868591       26209  ...        24              4
13    842603       26209  ...        24              4
14   1765563       26209  ...        24              4
15   1730767       26209  ...        24              4
16    672740       26209  ...        24              4
17    459187  
```

<!-- grade_id: cell-0c48a0352b6596f5, points: 1, grade: true, solution: false -->
```python
# Task 5 - Test

assert isinstance(final, pd.DataFrame), 'Have you created a DataFrame named `final`?'

```

**Task 6** - Top five products - value counts

Note that every record in `final` is a single instance in which someone ordered one of the top five products.  Your task is to calcluate the number of times each product was ordered.

* Calculate the **number of times** each product was ordered using `value_counts()`. Name the result `food_frequency`.
* Assign the most frequently ordered food to the variable `top_item` (your variable should be a string; make sure to use the exact formatting returned above).

<!-- grade_id: cell-9a3735ded06b8419, points: 0, grade: false, solution: true -->
```python
# Task 6

# Determine the number of times each product was ordered

# YOUR CODE HERE
food_frequency = final["product_name"].value_counts()

# View your results
food_frequency.head()
#food_frequency.shape
```

Output:
```
Bag of Organic Bananas    561
Organic Baby Spinach      360
Large Lemon               283
Organic Avocado           269
Limes                     197
Name: product_name, dtype: int64
```

```python
top_item = food_frequency.index[0]
top_item
```

Output:
```
'Bag of Organic Bananas'
```

<!-- grade_id: cell-e7858bb55bf086fb, points: 1, grade: true, solution: false -->
```python
# Task 6 - Test
assert food_frequency.shape == (5,), 'Check to make sure you merged correctly'

```

**Task 7** - Using `groupby()` and `mean()`

In this next step, we're going to import one last CSV file and look at the time of day an order was made. Some of the code has been provided for you - your task is to complete the merging and look at the mean time of the top five orders.

**Completed for you**
* Imported `orders_sm.csv` as the DataFrame `orders`.
* Merged `orders` and `final` into a DataFrame called `final_times` that only includes the top five foods.

**Your tasks**
* Use the `groupby()` and `mean()` functions to calculate the mean time each product was ordered (`order_hour_of_day`).
* Name the results of your calculation `mean_hours`.

<!-- grade_id: cell-8932129638b4d019, points: 0, grade: false, solution: true -->
```python
# Task 7

# COMPLETED FOR YOU
# Import orders.csv and name the DataFrame orders
data_url4 = 'https://raw.githubusercontent.com/LambdaSchool/data-science-practice-datasets/main/unit_1/Instacart_small/orders_sm.csv'
orders = pd.read_csv(data_url4)
# Merge orders and final into the DataFrame final_time
final_time = pd.merge(orders, final, on='order_id', how='right')

# Use groupby() and mean() to calculate the mean time each product was ordered

# YOUR CODE HERE
mean_hours = final_time.groupby('product_name')['order_hour_of_day'].mean()

# View your results
mean_hours.head()
```

Output:
```
product_name
Bag of Organic Bananas    13.200000
Large Lemon               16.000000
Limes                     10.333333
Organic Avocado           17.500000
Organic Baby Spinach      14.500000
Name: order_hour_of_day, dtype: float64
```

<!-- grade_id: cell-22190fff81957bb7, points: 1, grade: true, solution: false -->
```python
# Task 7 - Test
assert mean_hours.shape == (5,), 'Check to make sure you merged correctly'

```

**Task 8** - Mean order time

Using the results from above, determine the mean order time for two of the top five items: "Limes" and "Organic Baby Spinach"

* Find the mean order time for "Limes" and assign it to the variable `time_limes`. Your answer should be a float defined to one decimal place.
* Find the mean order time for "Organic Baby Spinach" and assign it to the variable `time_spinach`. You answer should be a float defined to one decimal place.

<!-- grade_id: cell-b160e1511b2388db, points: 0, grade: false, solution: true -->
```python
# Task 8

# YOUR CODE HERE

time_limes = float ("{:.1f}".format(mean_hours.loc["Limes"]))
time_spinach = float ("{:.1f}".format(mean_hours.loc["Organic Baby Spinach"]))
print (time_limes, time_spinach)
```

Output:
```
10.3 14.5
```

**Task 8 Test**

<!-- grade_id: cell-c138ab110b2bcf90, points: 1, grade: true, solution: false -->
```python
# Task 8 - Test

```