## Autograded Notebook (Canvas & CodeGrade)

This notebook will be automatically graded. It is designed to test your answers and award points for the correct answers. Following the instructions for each Task carefully.

### Instructions

* **Download this notebook** as you would any other ipynb file
* **Upload** to Google Colab or work locally (if you have that set-up)
* **Delete `raise NotImplementedError()`**
* Write your code in the `# YOUR CODE HERE` space
* **Execute** the Test cells that contain `assert` statements - these help you check your work (others contain hidden tests that will be checked when you submit through Canvas)
* **Save** your notebook when you are finished
* **Download** as a `ipynb` file (if working in Colab)
* **Upload** your complete notebook to Canvas (there will be additional instructions in Slack and/or Canvas)

# Lambda School Data Science - Unit 1 Sprint 1 Module 4

## Make Explanatory Visualizations

### Module Learning Objectives

- Identify the appropriate visualization type for a particular variable type and research question 
- Use Matplotlib to visualize distributions and relationships with continuous and discrete variables
- Add emphasis and annotations to transform visualizations from exploratory to explanatory
- Remove clutter from visualizations
- Identify misleading visualizations and how to fix them

### Total Points: 5

### Introduction

The first few modules in this Sprint have focused on exploring datasets, manipulating DataFrames, and creating new features. Now we're going to create some visualizations!

### Dataset Description

Researchers recorded data on sleep duration as well as a set of ecological and constitutional variables for a selection of mammal species.

The data dictionary can be accessed here: https://github.com/LambdaSchool/data-science-practice-datasets/tree/main/unit_1/Sleep in the READ ME file.

*Source: Allison, T. and Cicchetti, D. (1976), "Sleep in Mammals: Ecological and Constitutional Correlates", Science, November 12, vol. 194, pp. 732-734.*


**Task 1** -  Import the `sleep.csv` file and load it as a DataFrame named `Sleep`.
* Don't forget to include any import statements you need to create a DataFrame from a .csv file
* The `sleep.csv` file can be accessed using `sleep_url` which is provided for you
* Load the `sleep.csv` file and as a DataFrame named `Sleep`
* Print the first 5 rows of `Sleep`

<!-- grade_id: cell-565969801f9af810, points: 0, grade: false, solution: true -->
```python
# Task 1 

# Access sleep.csv with this url
sleep_url = 'https://raw.githubusercontent.com/LambdaSchool/data-science-practice-datasets/main/unit_1/Sleep/Sleep.csv'

# Read in the sleep.csv file included above as a DataFrame named Sleep and print the first 5 rows.
# Don't forget any necessary import statements.

# YOUR CODE HERE
raise NotImplementedError()

# View the DataFrame
Sleep.head()
```

**Task 1 Test**

<!-- grade_id: cell-22e1084a89e39897, points: 1, grade: true, solution: false -->
```python
# Task 1 - Test

assert isinstance(Sleep, pd.DataFrame), 'Have you created a DataFrame named `Sleep` (check your capitalization)?'
assert len(Sleep) == 42

```

**Task 2** - Plot a histogram of `Parasleep`, which is the number of hours of dreaming sleep each mammal slept during a 24-hour period.
* The import statement to `import matplotlib.pyplot as plt` is included for you.
* Use the template below and **replace the #### as specified in the instructions**
* Plot a histogram of `Parasleep` from the `Sleep` DataFrame
* The x-axis label should read `'Total hours of dreaming sleep'`
* The y-label should read `'Frequency`'
* The title should read `'Daily dreaming sleep in mammal species'`

**Note:** UNCOMMENT code lines to complete the plotting task

```python
# Task 2

# UNCOMMENT the code lines to complete the Task

# Import matplotlib
import matplotlib.pyplot as plt

fig, ax = plt.subplots()

# Plot a histogram of Parasleep from the Sleep DataFrame
#ax.hist(####)

# Specify the axis labels and plot title
#ax.set_xlabel(####) 
#ax.set_####('Frequency') 
####_title('Daily dreaming sleep in mammal species') 

plt.show()
```

<!-- grade_id: cell-a99053d0685466ef, points: 0, grade: false, solution: true -->
```python
# Task 2 SOLUTION 

# IGNORE the YOUR CODE HERE - your code is completed above
# You will see this solution when you submit the notebook

# YOUR CODE HERE
raise NotImplementedError()
```

**Task 3** - Prepare `Sleep` data for plotting

* Create a subset of `Sleep` that includes only records of mammals in `Danger` category 1.  Name this DataFrame `Danger_1`
* Create a subset of `Sleep` that includes only records of mammals in `Danger` category 5.  Name this DataFrame `Danger_5`

<!-- grade_id: cell-337cd8cb51d5f637, points: 0, grade: false, solution: true -->
```python
# Task 3

# Set up danger category 1 and danger category 5

# YOUR CODE HERE
raise NotImplementedError()

# View your DataFrames
print(Danger_1.head())
print(Danger_5.head())
```

**Task 3 Test**

<!-- grade_id: cell-c3505d0306800048, points: 1, grade: true, solution: false -->
```python
# Task 3 - Test

assert isinstance(Danger_1, pd.DataFrame), 'Have you created a DataFrame named Danger_1 (check your capitalization)?'
assert isinstance(Danger_5, pd.DataFrame), 'Have you created a DataFrame named Danger_5 (check your capitalization)?'

```

**Task 4** - Plot side-by-side box plots of `Parasleep` to compare the distribution of `Parasleep` for mammals in `Danger` category 1 and `Danger` category 5.

Use the template below and **replace the #### as specified in the instructions**

* Plot side-by-side boxplots of `Parasleep` from the `Sleep` DataFrame that shows the distribution of `Parasleep` separately for mammals in `Danger` categories 1 and 5.
* Label the boxplots `'Least danger'` and `'Most danger'`
* The x-axis label should read `'Total hours of dreaming sleep'`
* The y-label should read `'Danger category`'
* The title should read `'Daily dreaming sleep in mammal species'`
* The `vert` parameter in the boxplot should be set to `False`.

**The plotting code will not be autograded** but it is still required for completing the project.

```python
# Task 4 - Plotting

# UNCOMMENT the code lines to complete the Task

fig, ax = plt.subplots()

# Plot the side-by-side boxplots           
#ax.boxplot([ #### , Danger_5['Parasleep']], labels=['Least danger','####'], ####=False)

# Label the figure
#ax.set_xlabel('Hours of dreaming sleep')
#ax.set_ylabel('Danger category')
#ax.set_title('Dreaming sleep for mammals that experience low and high levels of danger')

plt.show()
```

<!-- grade_id: cell-5ba7c2434d312bda, points: 0, grade: false, solution: true -->
```python
# Task 4 SOLUTION

# IGNORE the YOUR CODE HERE - your code is completed above
# You will see this solution when you submit the notebook

# YOUR CODE HERE
raise NotImplementedError()
```

**Task 5** - Sort the `Sleep` DataFrame

* Sort the `Sleep` DataFrame by `Gest`.  Name the sorted DataFrame `Sleep_sorted`

<!-- grade_id: cell-c0ab45b6dbe40522, points: 0, grade: false, solution: true -->
```python
# Task 5

# Sort Sleep by Gest

# YOUR CODE HERE
raise NotImplementedError()

# View the results
Sleep_sorted.head()
```

**Task 5 - Test**

<!-- grade_id: cell-7420f399f5d1cef9, points: 1, grade: true, solution: false -->
```python
# Task 5 - Test

assert Sleep_sorted.iloc[0, 0] == 'N_American_opossum', 'Double-check your DataFrame sorting'

### No hidden tests
```

**Task 6** - Plot a line plot of `Parasleep` by `Gest`, the gestational time for each mammal.

Use the template below and **replace the #### as specified in the instructions**

* Plot a line plot using the `Sleep_sorted` DataFrame with each mammal's value of `Gest` on the x-axis and each mammal's value of `Parasleep` on the y-axis.
* In the `plot` statment, specify `'o'` for the marker, `'dashdot'` for the linestyle and `'b'` for the color.
* Label and title the graph using the statements provided for you.

**The plotting code will not be autograded** but it is still required for completing the project.

```python
# Task 6 - Plotting

# UNCOMMENT the code lines to complete the Task

fig, ax = plt.subplots()

#ax.plot(####, 
#    ####, 
#    marker=####,
#    linestyle=####,
#    color=####) 

#ax.set_xlabel('Time (month)') 
#ax.set_ylabel('Temperature (degrees F)') 
#ax.set_title('Monthly High Temperature in Seattle (degrees F)') 

plt.show()
```

<!-- grade_id: cell-eed45d8b40814c93, points: 0, grade: false, solution: true -->
```python
# Task 6 SOLUTION

# IGNORE the YOUR CODE HERE - your code is completed above
# You will see this solution when you submit the notebook

# YOUR CODE HERE
raise NotImplementedError()
```

**Task 7** - Create a normalized DataFrame

* Use `value_counts` with `normalize=True` to calculate the proportion of mammals in the `Sleep` DataFrame that are in each `Danger` category.  Save your results in a DataFrame named `Danger_prop`

<!-- grade_id: cell-2990d789c1d038b8, points: 0, grade: false, solution: true -->
```python
# Task 7

# YOUR CODE HERE
raise NotImplementedError()

# View the DataFrame
Danger_pct.head()
```

```python
Danger_pct.iloc[0,0]
```

**Task 7 - Test**

<!-- grade_id: cell-7833404c372083ae, points: 1, grade: true, solution: false -->
```python
# Task 7 - Test

assert isinstance(Danger_prop, pd.DataFrame), 'Have you created a DataFrame named Danger_prop?'
assert isinstance(Danger_pct, pd.DataFrame), 'Have you created a DataFrame named Danger_pct?'

```

**Task 8** - Plot the percent of animals in each `danger` category using a pie chart.

* Plot a pie chart using the `Danger` variable from the `Danger_pct` DataFrame.  
* Set the labels of the plot using `Danger_pct.index`
* Some other graphical parameters as well as the plot title have been set for you.

**The plotting code will not be autograded** but it is still required for completing the project.

```python
# Task 8 - Plotting

# UNCOMMENT the code lines to complete the Task

# Create the pie chart

fig, ax = plt.subplots()

#ax.pie(####, labels=####, autopct='%1.1f%%', startangle=90)
#ax.set_title('Percent of mammals in each danger category')

plt.show()
```

<!-- grade_id: cell-ade272759a14cda0, points: 0, grade: false, solution: true -->
```python
# Task 8 SOLUTION

# IGNORE the YOUR CODE HERE - your code is completed above
# You will see this solution when you submit the notebook

# YOUR CODE HERE
raise NotImplementedError()
```

**Task 9** - Create a new feature

* Use `.loc` to create a new feature in `Sleep` called `Short life` that takes on the values: 
    * `Less than 30 years` if the mammal's lifespan is less than 30 years long
    * `30 or more years` if the mammal's lifespan is 30 years or longer
* Use `.value_counts()` to calculate the frequency of `Short life`.  Save your results to a DataFrame named `Life_counts`

<!-- grade_id: cell-eea76e97d9e1d059, points: 0, grade: false, solution: true -->
```python
# Task 9

# YOUR CODE HERE
raise NotImplementedError()

# View the results
Life_counts.head()
```

**Task 9 - Test**

<!-- grade_id: cell-ce26aa435b114a62, points: 1, grade: true, solution: false -->
```python
# Task 9 - Test

assert 'Short life' in Sleep.columns, 'Have you created the new feature column?'
assert isinstance(Life_counts, pd.DataFrame), 'Have you created a DataFrame named Life_counts?'

```

**Task 10** - Plot the proportion of mammals in the `Sleep` dataset that had short (< 30 year) and long (> 30 year) lifespans.

* Plot a bar chart with the index of `Life_counts` as the bar labels and the values of the `Short life` column of `Life_counts` as the bar heights.
* Rotate the x tick labels 45 degrees.
* Other graphical parameters, axis labels, and the title have been set for you.

**The plotting code will not be autograded** but it is still required for completing the project.

```python
# Task 10 - Plotting

# UNCOMMENT the code lines to complete the Task

fig, ax = plt.subplots()

#ax.bar(Life_counts####,Life_counts####)

#ax.set_xlabel('Mammal lifespan') 
#ax.set_xticklabels(Life_counts.index, rotation=####) 

#ax.set_ylabel('Frequency') 
#ax.set_title('Number of Mammals with Long and Short Life Expectancies') 

plt.show()
```

<!-- grade_id: cell-485432aaba3752cd, points: 0, grade: false, solution: true -->
```python
# Task 10 - Plotting

# IGNORE the YOUR CODE HERE - your code is completed above
# You will see this solution when you submit the notebook

        
# YOUR CODE HERE
raise NotImplementedError()
```