## Autograded Notebook (Canvas & CodeGrade)

This notebook will be automatically graded. It is designed to test your answers and award points for the correct answers. Following the instructions for each Task carefully.

### Instructions

* **Download this notebook** as you would any other ipynb file
* **Upload** to Google Colab or work locally (if you have that set-up)
* **Delete `raise NotImplementedError()`**
* Write your code in the `# YOUR CODE HERE` space
* **Execute** the Test cells that contain `assert` statements - these help you check your work (others contain hidden tests that will be checked when you submit through Canvas)
* **Save** your notebook when you are finished
* **Download** as a `ipynb` file (if working in Colab)
* **Upload** your complete notebook to Canvas (there will be additional instructions in Slack and/or Canvas)

# Data Science - Unit 1 Sprint 1 Module 2

## Making Features

### Module Learning Objectives

- Understand the purpose of feature engineering
- Work with strings in pandas
- Work with dates and times in pandas
- Modify or create columns of a dataframe using the `.apply()` function

### Resources

- Python Data Science Handbook
  - [Chapter 3.10](https://jakevdp.github.io/PythonDataScienceHandbook/03.10-working-with-strings.html), Vectorized String Operations - eg. .strip(), .split(), .replace() and list comprehensions
  - [Chapter 3.11](https://jakevdp.github.io/PythonDataScienceHandbook/03.11-working-with-time-series.html), Working with Time Series - lots more details about time using and not using Pandas


### Notebook points: 11

### Introduction

This project is going to focus on learning how to create new features from exisiting data. There are many ways to create features in your datasets; we'll only focus on a few here. Let's get started!

**Task 1** - Load the data set

* Using the provided URL, load the Ames Housing Data into a DataFrame called `house`
* View your DataFrame

<!-- grade_id: cell-604a4cd13da263c1, points: 0, grade: false, solution: true -->
```python
# Task 1

# Imports
import pandas as pd

# Dataset URL

data_url = 'https://raw.githubusercontent.com/LambdaSchool/data-science-practice-datasets/main/unit_1/Housing/Ames%20Iowa%20Housing%20Data.csv'

# YOUR CODE HERE
house = pd.read_csv(data_url)

# View your DataFrame
print(house)
print()

print( house.isnull().sum()[ (house.isnull().sum() > 0) ] )
print()

#house.info()

```

Output:
```
Id  MSSubClass MSZoning  ...  SaleType  SaleCondition SalePrice
0        1          60       RL  ...        WD         Normal    208500
1        2          20       RL  ...        WD         Normal    181500
2        3          60       RL  ...        WD         Normal    223500
3        4          70       RL  ...        WD        Abnorml    140000
4        5          60       RL  ...        WD         Normal    250000
...    ...         ...      ...  ...       ...            ...       ...
1455  1456          60       RL  ...        WD         Normal    175000
1456  1457          20       RL  ...        WD         Normal    210000
1457  1458          70       RL  ...        WD         Normal    266500
1458  1459          20       RL  ...        WD         Normal    142125
1459  1460          20       RL  ...        WD         Normal    147500

[1460 rows x 81 columns]

LotFrontage      259
Alley           1369
MasVnrType         8
MasVnrArea         8
BsmtQual          37
BsmtCond    
```

**Task 1 Test**

<!-- grade_id: cell-4be6ce6f2ae0213b, points: 1, grade: true, solution: false -->
```python
# Task 1 - Test

# These tests are for you to check your work before submitting
assert isinstance(house, pd.DataFrame), 'Have you created a DataFrame named house?'

# Hidden tests - you'll see results when you submit to Canvas
```

**Task 2** - Create a new feature

Now that we have the data loaded, we're going to create a new feature. We can use the some of the existing features, in this case, the total number of bathrooms in one column.

* Create a new variable called `Total_Bathrooms` that contains the total number of full and half bathrooms in the house.  

*Hint: Identify all the columns with bath in the title and add their values together. There are four columns to add together.  Add one to the total for each full or half bathroom.*

<!-- grade_id: cell-98808d8ac7de9a36, points: 0, grade: false, solution: true -->
```python
# Task 2

# YOUR CODE HERE
house['Total_Bathrooms'] = house['FullBath'] + house['HalfBath']
print(house['Total_Bathrooms'])
```

Output:
```
0       3
1       2
2       3
3       1
4       3
       ..
1455    3
1456    2
1457    2
1458    1
1459    2
Name: Total_Bathrooms, Length: 1460, dtype: int64
```

**Task 2 Test**

<!-- grade_id: cell-f9c97fbd0846aa11, points: 1, grade: true, solution: false -->
```python
# Task 2 - Test

assert 'Total_Bathrooms' in house.columns, 'Did you add the new column?'

# Hidden tests - you'll see results when you submit to Canvas

# View all the bathroom columns
house[['Total_Bathrooms','BsmtFullBath','BsmtHalfBath','FullBath','HalfBath']].head()
```

Output:
```
Total_Bathrooms  BsmtFullBath  BsmtHalfBath  FullBath  HalfBath
0                3             1             0         2         1
1                2             0             1         2         0
2                3             1             0         2         1
3                1             1             0         1         0
4                3             1             0         2         1
```

**Task 3** - Create a new feature using encoding

Next we're going to create another new column by "encoding" an exisitng column. Using the `LotArea` column, we'll assign values according the the area of the lot.

* Create a new column called `Large_Lot` that takes on the following values:
    * 0 when `LotArea` is **less** than 10,000
    * 1 when `LotArea` is **greater than or equal** to 10,000

<!-- grade_id: cell-baf76e10fc75872a, points: 0, grade: false, solution: true -->
```python
# Task 3

# YOUR CODE HERE
house.loc[house['LotArea'] < 10000, "Large_Lot"] = 0
house.loc[house['LotArea'] > 10000, "Large_Lot"] = 1

# View the new column
house[['LotArea', 'Large_Lot']]
```

Output:
```
LotArea  Large_Lot
0        8450        0.0
1        9600        0.0
2       11250        1.0
3        9550        0.0
4       14260        1.0
...       ...        ...
1455     7917        0.0
1456    13175        1.0
1457     9042        0.0
1458     9717        0.0
1459     9937        0.0

[1460 rows x 2 columns]
```

**Task 3 Test**

<!-- grade_id: cell-dc861d3d6dbf7bec, points: 1, grade: true, solution: false -->
```python
# Task 3 - Test

assert 'Large_Lot' in house.columns, 'Did you add the Large_Lot column?'

# Hidden tests - you'll see results when you submit to Canvas
```

**Task 4** - Create another feature

Let's continue with creating new features. This time, we'll focus on encoding for the sale month.

* Create a new column called `Summer_Sale` that takes on the following values:
    * 0 if the sale month was in September - May (`ModSold` = 1-5, 9-12)
    * 1 if the sale month was in June, July or August (`MoSold` = 6, 7 or 8)

<!-- grade_id: cell-077200eae8a405e2, points: 0, grade: false, solution: true -->
```python
# Task 4

# YOUR CODE HERE
house['Summer_Sale'] = 1
house.loc[(house['MoSold'] < 6) | (house['MoSold'] > 8), "Summer_Sale"] = 0


# View the new column
#print(house[['MoSold', 'Summer_Sale']])


#All DataFrame Rows meeting both summer conditions
summer_sales = house.loc[(house['MoSold'] > 5) & (house['MoSold'] < 9)]

#Only the month columns for that DataFrame
print(summer_sales[['MoSold', 'Summer_Sale']])
```

Output:
```
MoSold  Summer_Sale
6          8            1
11         7            1
13         8            1
15         7            1
18         6            1
...      ...          ...
1435       7            1
1449       8            1
1453       7            1
1455       8            1
1459       6            1

[609 rows x 2 columns]
```

**Task 4 Test**

<!-- grade_id: cell-7df13700c084fe99, points: 1, grade: true, solution: false -->
```python
# Task 4 - Test

assert 'Summer_Sale' in house.columns, 'Did you add the Summer_Sale column?'

# Hidden tests - you'll see results when you submit to Canvas
```

**Task 5**

Now we'll revisit the LendingClub data. The statements to import the `loans` DataFrame have been provided for you.  Run the code block below without changing anything to load the `loans` data.

* Load the data into a DataFrame called `loans`.
* Make sure to view the `loans` DataFrame to see what data we're working with.

<!-- grade_id: cell-a03fb13afedb44b5, points: 0, grade: false, solution: true -->
```python
# Task 5

# Dataset URL

# Dataset URL
loans_data = 'https://raw.githubusercontent.com/LambdaSchool/data-science-practice-datasets/main/unit_1/LendingClub/LoanStats_2018Q4_sm.csv'

# YOUR CODE HERE
loans = pd.read_csv(loans_data)

# View the dataset!
print(loans['earliest_cr_line'])
```

Output:
```
0        Oct-1998
1        Sep-2015
2        Jun-2003
3        Feb-2005
4        Feb-2008
           ...   
29995    Jan-2000
29996    Sep-1985
29997    Dec-2007
29998    Apr-2008
29999    Apr-1991
Name: earliest_cr_line, Length: 30000, dtype: object
```

**Tast 5 Test**

<!-- grade_id: cell-365203eb4c6651fd, points: 1, grade: true, solution: false -->
```python
# Task 5 - Test

assert isinstance(loans, pd.DataFrame), 'Have you created a DataFrame named loans?'
assert loans.shape == (30000, 144), 'Double check your DataFrame size.'

# NO hidden tests for this task
```

**Introduction** - Approaching a problem two different ways.

In the Guided Project, we learned how to find the earliest credit year using the built-in date-time format.  However, there is often more than one way to come up with the same solution.  

In the following questions, we will work through the steps to create new variables called `Earliest_Credit_Year` and `issue_year` where we use the `.split()` function to split on the month and year in a string variable, and return just the year value. We can then use those variables to calculate the length of credit history in years.  

If we do everything correctly, we should calculate *about* the same length of longest credit history that we did in class working with the date-time format.  The answers won't exactly match because we are making different assumptions about when in each month and year the loans were taken out, but it will give you a flavor for different ways of approaching the same problem.

**Task 6** - Create a simple test case

The `credit` variable has been created for you with the value 'Jun-1979'

* Use the `.split('-')` function to separate the month and year parts of the data.
* Name the results of the .split('-') function `fields`: 
    * assign `fields[0]` to the variable `month`
    * assign `fields[1]` to the variable `year`

<!-- grade_id: cell-2e0e85d026ef0577, points: 0, grade: false, solution: true -->
```python
# Task 6

# Don't change or delete
credit = 'Jun-1979' 

# YOUR CODE HERE
fields = credit.split('-')
month = fields[0]
year = fields[1]

# Look at the year variable both as a string and as a float
print('String: ', year)
print('Float: ', float(year))
```

Output:
```
String:  1979
Float:  1979.0
```

**Task 6 Test**

<!-- grade_id: cell-95ad43446cc5faaa, points: 1, grade: true, solution: false -->
```python
# Task 6 - Test

# Hidden tests - you'll see results when you submit to Canvas
```

**Task 7** - Create a function

Use your answer from **Task 6** to write a function called `credit_yr` that takes in the contents of a cell that is formatted as "month-year" and returns the year.

* Complete the definition of the function `credit_yr` below:
    * one argument `cell_contents` as input (provided for you)
    * splits the `cell_content` on a `-` and calls the result `fields`
    * creates the variable `year` which has the value `fields[1]`
    * returns `year` as a float variable
* Run your function using the simple test case in Task 6 (`credit='Jun-1979'`). Assign the output of the function to the variable `year_function`.

<!-- grade_id: cell-67705c48e95b83db, points: 0, grade: false, solution: true -->
```python
# Task 7

# Complete the function definition
def credit_yr(cell_contents):
  fields = cell_contents.split('-')
  year = fields[1]
  return float(year)

# Print your year_function variable
year_function = credit_yr(credit)
print(year_function)
```

Output:
```
1979.0
```

**Task 7 Test**

<!-- grade_id: cell-e850dcb0651cdf47, points: 1, grade: true, solution: false -->
```python
# Task 7 - Test

assert isinstance(year_function, float), 'Make sure your year variable is a float!'

# Hidden tests - you'll see results when you submit to Canvas
```

**Task 8** - Apply our function

Now we're going to use the function we created and apply it to a column in our `loans` DataFrame. 

* Use the `.apply()` function to apply the function to every cell in the `earliest_cr_line` column.
* Assign the results to a new column called `Earliest_Credit_Year`.
* View the top five rows of the `Earliest_Credit_Year` and `earliest_cr_line` columns to make sure the variables were created correctly (code provided for you).

<!-- grade_id: cell-190721b22d999b13, points: 0, grade: false, solution: true -->
```python
# Task 8

loans = loans.copy()

# YOUR CODE HERE
loans['Earliest_Credit_Year'] = loans['earliest_cr_line'].apply(credit_yr)

# Print your columns
print(loans[['Earliest_Credit_Year','earliest_cr_line']])
```

Output:
```
Earliest_Credit_Year earliest_cr_line
0                    1998.0         Oct-1998
1                    2015.0         Sep-2015
2                    2003.0         Jun-2003
3                    2005.0         Feb-2005
4                    2008.0         Feb-2008
...                     ...              ...
29995                2000.0         Jan-2000
29996                1985.0         Sep-1985
29997                2007.0         Dec-2007
29998                2008.0         Apr-2008
29999                1991.0         Apr-1991

[30000 rows x 2 columns]
```

**Task 8 Test**

<!-- grade_id: cell-1c6181c54ca22ae6, points: 1, grade: true, solution: false -->
```python
# Task 8 - Test

assert 'Earliest_Credit_Year' in loans.columns, 'Did you add the "Earliest_Credit_Year" column?'
assert loans.shape == (30000, 145), 'Double check your DataFrame size.'

# Hidden tests - you'll see results when you submit to Canvas
```

**Task 9** - Apply the function to a new column

So next, we're going to use the `credit_yr` function to create a new variable called `Issue_Year` from the `issue_d` variable. Follow the sames steps as you did in Task 8.

* Apply `credit_yr` to the `issue_d` column.
* Name your new column `Issue_Year`.
* View the `Issue_Year`,`issue_d` columns (code provided for you).

<!-- grade_id: cell-b31b3dfe0d069095, points: 0, grade: false, solution: true -->
```python
# Task 9

# YOUR CODE HERE
loans['Issue_Year'] = loans['issue_d'].apply(credit_yr)

# Print your columns
print(loans[['Issue_Year','issue_d']])
```

Output:
```
Issue_Year   issue_d
0          2018.0  Dec-2018
1          2018.0  Oct-2018
2          2018.0  Oct-2018
3          2018.0  Nov-2018
4          2018.0  Dec-2018
...           ...       ...
29995      2018.0  Dec-2018
29996      2018.0  Oct-2018
29997      2018.0  Oct-2018
29998      2018.0  Oct-2018
29999      2018.0  Dec-2018

[30000 rows x 2 columns]
```

**Task 9**

<!-- grade_id: cell-1ae1a8aef5f84b70, points: 1, grade: true, solution: false -->
```python
# Task 9 - Test

assert 'Issue_Year' in loans.columns, 'Did you add the "Issue_Year" column?'
assert loans.shape == (30000, 146), 'Double check your DataFrame size.'

# Hidden tests - you'll see results when you submit to Canvas
```

**Task 10** - Calculate length of time from your new columns

Now we're going to use the two new columns you created to calculate the length of credit history in years and in days.

* Create a new column called `Credit_History_Years` which is the difference between `Issue_Year` and `Earliest_Credit_Year`
* Create another new column called `Credit_History_Days` which is the value in the column you created above expect multipled by 365.25 to convert years to days.
* View the top five rows of the `loans` DataFrame (code has been provided).

Hint: If you are getting an error message, make sure to check the data type of the `year` values output by your function.

<!-- grade_id: cell-2f411113665bb37a, points: 0, grade: false, solution: true -->
```python
# Task 10

# YOUR CODE HERE
loans['Credit_History_Years'] = loans['Issue_Year'] - loans['Earliest_Credit_Year']
loans['Credit_History_Days'] = loans['Credit_History_Years'] * 365.25

display(loans[['Issue_Year','Earliest_Credit_Year','Credit_History_Years','Credit_History_Days']])
```

Output:
```
Issue_Year  ...  Credit_History_Days
0          2018.0  ...              7305.00
1          2018.0  ...              1095.75
2          2018.0  ...              5478.75
3          2018.0  ...              4748.25
4          2018.0  ...              3652.50
...           ...  ...                  ...
29995      2018.0  ...              6574.50
29996      2018.0  ...             12053.25
29997      2018.0  ...              4017.75
29998      2018.0  ...              3652.50
29999      2018.0  ...              9861.75

[30000 rows x 4 columns]
```

**Task 10 Test**

<!-- grade_id: cell-e3694613b297ae85, points: 1, grade: true, solution: false -->
```python
# Task 10 - Test

assert 'Credit_History_Years' in loans.columns, 'Did you add the correct column?'
assert 'Credit_History_Days' in loans.columns, 'Did you add the correct column converted to days?'
assert loans.shape == (30000, 148), 'Double check your DataFrame size.'

# Hidden tests - you'll see results when you submit to Canvas
```

**Task 11** - Calculate the maximum credit length

Finally, we're going to use results from the `Credit_History_Years` column and compare it to we got in the Guided Project.

* Find the maximum values in the `Credit_History_Years` column and assign it to the variable `max_credit`. Make sure you have a float value defined out to one decimal place.
* Compare the the value from the Guided Project (refer to your notebook from class).

Note - it won't be *exactly* the same because both methods are working in different ways, but they will give you a "pretty close" answer.

<!-- grade_id: cell-421ef18e66bbaab0, points: 0, grade: false, solution: true -->
```python
# Task 11

# YOUR CODE HERE
max_credit = loans['Credit_History_Years'].max()
print(max_credit)

max_credit_index = loans['Credit_History_Years'].sort_values(ascending=False).index[0]

print(loans.iloc[max_credit_index][['Issue_Year','Earliest_Credit_Year','Credit_History_Years','Credit_History_Days']])
```

Output:
```
68.0
Issue_Year               2018
Earliest_Credit_Year     1950
Credit_History_Years       68
Credit_History_Days     24837
Name: 1635, dtype: object
```

**Task 11 Test**

<!-- grade_id: cell-b39de9ca5d93a343, points: 1, grade: true, solution: false -->
```python
# Task 11 - Test

# Hidden tests - you'll see results when you submit to Canvas

```