## Autograded Notebook (Canvas & CodeGrade)

This notebook will be automatically graded. It is designed to test your answers and award points for the correct answers. Following the instructions for each Task carefully.

### Instructions

* **Download this notebook** as you would any other ipynb file
* **Upload** to Google Colab or work locally (if you have that set-up)
* **Delete `raise NotImplementedError()`**
* Write your code in the `# YOUR CODE HERE` space
* **Execute** the Test cells that contain `assert` statements - these help you check your work (others contain hidden tests that will be checked when you submit through Canvas)
* **Save** your notebook when you are finished
* **Download** as a `ipynb` file (if working in Colab)
* **Upload** your complete notebook to Canvas (there will be additional instructions in Slack and/or Canvas)

# Lambda School Data Science - Unit 1 Sprint 1 Module 3

## Join and Reshape Data 

### Module Learning Objectives

* Concatenate data using the pandas concat method
* Merge data using pandas merge
* Define the concept of tidy data and describe the format
* Tidy data

### Helpful Resources

* [Pandas Cheat Sheet](https://github.com/pandas-dev/pandas/blob/master/doc/cheatsheet/Pandas_Cheat_Sheet.pdf)
* Python Data Science Handbook
    * [Chapter 3.6](https://jakevdp.github.io/PythonDataScienceHandbook/03.06-concat-and-append.html), Combining Datasets: Concat and Append
    * [Chapter 3.7](https://jakevdp.github.io/PythonDataScienceHandbook/03.07-merge-and-join.html), Combining Datasets: Merge and Join
    * [Chapter 3.8](https://jakevdp.github.io/PythonDataScienceHandbook/03.08-aggregation-and-grouping.html), Aggregation and Grouping
    * [Chapter 3.9](https://jakevdp.github.io/PythonDataScienceHandbook/03.09-pivot-tables.html), Pivot Tables

## Introduction

In this Module Project, we're going to focus on manipulating data: loading multiple datasets from different locations, concatenating, merging, and even the fancy pandas `groupby` method.

## Dataset Description

You can view the CSV files [here](https://github.com/LambdaSchool/data-science-practice-datasets/tree/main/unit_1/Instacart_small), before you load them into your notebook.

**Task 1** - Load CSV files

We're going to look at the Instacart data and load the files as we need them. The Pandas and numpy import statements have been provided for you as well as the URL for the dataset.

* Using the provided URL (`data_url1`) load the product CSV file into a DataFrame called `product_names`.
* Print the top five records (or rows) of `product_names`.

<!-- grade_id: cell-ef51cab74649e553, points: 0, grade: false, solution: true -->
```python
# Task 1

# Imports 
import pandas as pd
import numpy as np

# Dataset URL
data_url1 = 'https://raw.githubusercontent.com/LambdaSchool/data-science-practice-datasets/main/unit_1/Instacart_small/product_names_sm.csv'
# YOUR CODE HERE
raise NotImplementedError()

# View the DataFrame
product_names.head()
```

<!-- grade_id: cell-cf00b9647815499e, points: 1, grade: true, solution: false -->
```python
# Task 1 - Test

assert isinstance(product_names, pd.DataFrame), 'Have you created a DataFrame named `product_names`?'
assert len(product_names) == 25000

```

**Task 2** - DataFrame concatenation

Next, we're going to load two additional CSV files. The URLs have been provided so your task is to load the files into the appropriately named DataFrames and then to concatenate them into one DataFrame.

* Load the CSV file at `data_url2` into a DataFrame called `order_products_prior`.
* Load the CSV file at `data_url3` into a DataFrame called `order_products_train`.
* Concatenate `order_products_prior` and `order_products_train` into one DataFrame called `order_products`.

<!-- grade_id: cell-0f442142d6bc8c39, points: 0, grade: false, solution: true -->
```python
# Task 2

# Dataset URLs
data_url2 = 'https://raw.githubusercontent.com/LambdaSchool/data-science-practice-datasets/main/unit_1/Instacart_small/order_products_prior_sm.csv'
data_url3 = 'https://raw.githubusercontent.com/LambdaSchool/data-science-practice-datasets/main/unit_1/Instacart_small/order_products_train_sm.csv'

# YOUR CODE HERE
raise NotImplementedError()

# Print the head to view your results
order_products.head()
```

<!-- grade_id: cell-3637e4ba19b79349, points: 1, grade: true, solution: false -->
```python
# Task 2 - Test

assert isinstance(order_products, pd.DataFrame), 'Have you created a DataFrame named `order_products`?'
assert len(order_products) == 50000

```

**Task 3** - Top 10 products

Execute the code below to identify the top 10 most ordered products. The `ten_products` variable has been created for you. 

* Use `type()` to determine the data type of the variable `ten_products`; name the result `data_type`.

<!-- grade_id: cell-ad8e4664b33be149, points: 0, grade: false, solution: true -->
```python
# Task 3

# Top 10 products: ten_products

ten_products = ["Banana",
                "Bag of Organic Bananas",
                "Organic Strawberries",
                "Organic Baby Spinach",
                "Organic Hass Avocado",
                "Organic Avocado",
                "Large Lemon",
                "Strawberries",
                "Limes",
                "Organic Whole Milk"]

# Determine the data type of ten_products

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-0e8a931d7088b469, points: 1, grade: true, solution: false -->
```python
# Task 3 - Test

```

**Task 4** - DataFrame subset using merge

Create a subset of the `product_names` dataset that includes just the top 10 most ordered products. The `ten_products` list has been converted to a DataFrame with the name `ten`. Your task is to merge these DataFrames.

* Merge `product_names` and `ten` so the resulting DataFrame only includes the 10 most ordered products.  Name that DataFrame `top_products`.

<!-- grade_id: cell-7818f18ec6921fb0, points: 0, grade: false, solution: true -->
```python
# Task 4

# Convert ten_products to a DataFrame called ten
ten = pd.DataFrame(ten_products, columns =['product_name']) 

# Create a subset of product_names with only the top 10 products.  
# Use your merge statement to include only the top 10 products

# YOUR CODE HERE
raise NotImplementedError()

# Print your DataFrame to view the merge results
top_products.head()
```

<!-- grade_id: cell-e329e3cdc129ba49, points: 1, grade: true, solution: false -->
```python
# Task 4 - Test
assert isinstance(top_products, pd.DataFrame), 'Have you created a DataFrame named `top_products`?'

```

**Task 5** - More DataFrame merging

Now we're going to merge the first DataFrame we created (`order_products`) with the one we just made above (`top_products`). This will result in a DataFrame with only the top 10 most ordered products.

* Merge `order_products` and `top_products` and name the resulting DataFrame `final`.
* Print the top 20 rows of `final` to view your result.

<!-- grade_id: cell-977f302dd89af1df, points: 0, grade: false, solution: true -->
```python
# Task 5

# Merge order_products and top_products 

# YOUR CODE HERE
raise NotImplementedError()

# Print your final DataFrame

```

<!-- grade_id: cell-0c48a0352b6596f5, points: 1, grade: true, solution: false -->
```python
# Task 5 - Test
assert isinstance(final, pd.DataFrame), 'Have you created a DataFrame named `final`?'

```

**Task 6** - Top 10 products value counts

Note that every record in `final` is a single instance in which someone ordered one of the top 10 products.  Your task is to calcluate the number of times each product was ordered.

* Calculate the number of times each product was ordered using `value_counts()`. Name the result `food_frequency`.
* Assign the most frequently ordered food to the variable `top_item` (your variable should be a string; make sure to use the exact formatting returned above).

<!-- grade_id: cell-9a3735ded06b8419, points: 0, grade: false, solution: true -->
```python
# Task 6

# Determine the number of times each product was ordered

# YOUR CODE HERE
raise NotImplementedError()

# View your results
food_frequency.head(10)
```

<!-- grade_id: cell-e7858bb55bf086fb, points: 1, grade: true, solution: false -->
```python
# Task 6 - Test
assert food_frequency.shape == (10,), 'Check to make sure you merged correctly'

```

**Task 7** - Using `groupby()` and `mean()`

In this next step, we're going to import one last CSV file and look at the time of day an order was made. Some of the code has been provided for you - your task is to complete the merging and look at the mean time of the top 10 orders.

**Completed for you**
* Imported `orders_sm.csv` as the DataFrame `orders`.
* Merged `orders` and `final` into a DataFrame called `final_times` that only includes the top 10 foods.

**Your tasks**
* Use the `groupby()` and `mean()` functions to calculate the mean time each product was ordered (`order_hour_of_day`).
* Name the results of your calculation `mean_hours`.

Note: You might notice some values in `mean_hours` have NaN for the times - this is because these items aren't represented in the smaller dataset we used in this Module Project.

<!-- grade_id: cell-8932129638b4d019, points: 0, grade: false, solution: true -->
```python
# Task 7

# COMPLETED FOR YOU
# Import orders.csv and name the DataFrame orders
data_url4 = 'https://raw.githubusercontent.com/LambdaSchool/data-science-practice-datasets/main/unit_1/Instacart_small/orders_sm.csv'
orders = pd.read_csv(data_url4)
# Merge orders and final into the DataFrame final_time
final_time = pd.merge(orders, final, on='order_id', how='right')

# Use groupby() and mean() to calculate the mean time each product was ordered

# YOUR CODE HERE
raise NotImplementedError()

# View your results
mean_hours.head(10)
```

<!-- grade_id: cell-22190fff81957bb7, points: 1, grade: true, solution: false -->
```python
# Task 7 - Test
assert mean_hours.shape == (10,), 'Check to make sure you merged correctly'

```

**Task 8** - Mean order time

Using the results from above, determine the mean order time for two of the top ten items: "Limes" and "Organic Baby Spinach"

* Find the mean order time for "Limes" and assign it to the variable `time_limes`. Your answer should be a float defined to one decimal place.
* Find the mean order time for "Organic Baby Spinach" and assign it to the variable `time_spinach`. You answer should be a float defined to one decimal place.

<!-- grade_id: cell-b160e1511b2388db, points: 0, grade: false, solution: true -->
```python
# Task 8

# YOUR CODE HERE
raise NotImplementedError()
```

**Task 8 Test**

<!-- grade_id: cell-c138ab110b2bcf90, points: 1, grade: true, solution: false -->
```python
# Task 8 - Test

```